---
title       : Insert the chapter title here
description : Insert the chapter description here
--- type:NormalExercise lang:r xp:100 skills:1 key:4a6d770497
## Test 2 


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
