---
title       : Insert the chapter title here
description : Insert the chapter description here
--- type:NormalExercise lang:r xp:100 skills:1 key:4a6d770497
## Test 2 


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```

--- type:VideoExercise lang:r xp:50 skills:1 key:b5a561bd28
## Testing out slides

*** =slidesKey
38f3a018444ffb15e2c18cf73c4a71be
